# Python for Data Science Bootcamp
## Mini Project 1

NOTE: This mini project has 4 problems.  Please write your code in the designated portion of the autograded answer cell.

For this problem set, we'll be using the Jupyter notebook and nbgrader:

![](jupyter.png)

©2020 Luis Morales. All rights reserved. 

## Problem 1 (1 point)

Given a list of integers, return indices of the two numbers such that they add up to a specific target. The indices should be in an increasing order. You may assume that each input would have exactly one solution, and you may not use the same element twice.

<!-- grade_id: twosum, points: 0, grade: false, solution: true -->
```python
def TwoSum(nums, target):
    """
    :type nums: List[int]
    :type target: int
    :rtype: List[int]
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: twosumtest, points: 1, grade: true, solution: false -->
```python
'''this TwoSum function should pass the following tests'''
assert (TwoSum([2, 7, 11, 15],9)==[0,1])
assert (TwoSum([3,2,4],6)==[1,2])
assert (TwoSum([0,4,3,0],0)==[0,3])
assert (TwoSum([-1,-2,-3,-4,-5],-8)==[2,4])
```

## Problem 2 (1 point)

Given a 32-bit signed integer, reverse digits of an integer. If the return number is outside this range [−2^31,2^31−1], then return 0.

example: input 12345 output 54321

<!-- grade_id: reverse, points: 0, grade: false, solution: true -->
```python
def myreverse(x):
    ''':type x: int
       :rtype: int
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: reversetest, points: 1, grade: true, solution: false -->
```python
assert(myreverse(120)==21)
assert(myreverse(123)==321)
assert(myreverse(-123)==-321)
assert(myreverse(1534236469)==0)
```

## Problem 3 (1 point)

Write a function to find the longest common prefix string amongst a list of strings. If there is no common prefix, return an empty string ''.

example: input ['what','where','who'] output 'wh'

<!-- grade_id: longestpre, points: 0, grade: false, solution: true -->
```python
def longestCommonPrefix(strs):
    """
    :type strs: List[str]
    :rtype: str
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: longestpretest, points: 1, grade: true, solution: false -->
```python
assert longestCommonPrefix(["flower","flow","flight"])=='fl'
assert longestCommonPrefix(['what','where','who'] )=='wh'
```

## Problem 4 (1 point)

There are 13 lamps on the table. Starting from 1, we turn off the lamp at every 5th, when the count reach 13, start from 1 again. Write a function that return the order of lamps being turned off. In this case, the return list would be [5, 10, 2, 8, 1, 9, 4, 13, 12, 3, 7, 11, 6]. Note that the turn-off lamps would not count again.

<!-- grade_id: lamp, points: 0, grade: false, solution: true -->
```python
def turn_off_order():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: lamptest, points: 1, grade: true, solution: false -->
```python
assert turn_off_order() == [5, 10, 2, 8, 1, 9, 4, 13, 12, 3, 7, 11, 6]
```